## Tips
- To avoid unpleasant surprises, I suggest you _run all cells in their order of appearance_ (__Cell__ $\rightarrow$ __Run All__).


- If the changes you've made to your solution don't seem to be showing up, try running __Kernel__ $\rightarrow$ __Restart & Run All__ from the menu.


- Before submitting your assignment, make sure everything runs as expected. First, restart the kernel (from the menu, select __Kernel__ $\rightarrow$ __Restart__) and then **run all cells** (from the menu, select __Cell__ $\rightarrow$ __Run All__).

## Reminder

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, UA email, and collaborators below:



Several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

```python
NAME = "Kathleen Costa"
# University of Arizona email address
EMAIL = "kathleencosta@arizona.edu"
# Names of any collaborators.  Write N/A if none.
COLLABORATORS = "N/A"
```

## Scratchpad

You are welcome to create new cells (see the __Cell__ menu) to experiment and debug your solution.

```python
%load_ext autoreload
%autoreload 2
```

# Mini Python tutorial

This course uses Python 3.11.

Below is a very basic (and incomplete) overview of the Python language... 

For those completely new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.11/library/stdtypes.html#common-sequence-operations).

```python
# This is a comment.  
# Any line starting with # will be interpreted as a comment

# this is a string assigned to a variable
greeting = "hello"

# If enclosed in triple quotes, strings can also be multiline:

"""
I'm a multiline
string.
"""

# let's use a for loop to print it letter by letter
for letter in greeting:
    print(letter)
    
# Did you notice the indentation there?  Whitespace matters in Python!

# here's a list of integers

numbers = [1, 2, 3, 4]

# let's add one to each number using a list comprehension
# and assign the result to a variable called res
# list comprehensions are used widely in Python (they're very Pythonic!)

res = [num + 1 for num in numbers]

# let's confirm that it worked
print(res)

# now let's try spicing things up using a conditional to filter out all values greater than or equal to 3...
print([num for num in res if not num >= 3])

# Python 3.7 introduced "f-strings" as a convenient way of formatting strings using templates
# For example ...
name = "Josuke"

print(f"{greeting}, {name}!")

# f-strings are f-ing convenient!


# let's look at defining functions in Python..

def greet(name):
    print(f"Howdy, {name}!")

# here's how we call it...

greet("partner")

# let's add a description of the function...

def greet(name):
    """
    Prints a greeting given some name.
    
    :param name: the name to be addressed in the greeting
    :type name: str
    
    """
    print(f"Howdy, {name}!")
    
# I encourage you to use docstrings!

# Python introduced support for optional type hints in v3.5.
# You can read more aobut this feature here: https://docs.python.org/3.7/library/typing.html
# let's give it a try...
def add_six(num: int) -> int:
    return num + 6

# this should print 13
print(add_six(7))

# Python also has "anonymous functions" (also known as "lambda" functions)
# take a look at the following code:

greet_alt = lambda name: print(f"Hi, {name}!")

greet_alt("Fred")

# lambda functions are often passed to other functions
# For example, they can be used to specify how a sequence should be sorted
# let's sort a list of pairs by their second element
pairs = [("bounce", 32), ("bighorn", 12), ("radical", 4), ("analysis", 7)]
# -1 is last thing in some sequence, -2 is the second to last thing in some seq, etc.
print(sorted(pairs, key=lambda pair: pair[-1]))

# we can sort it by the first element instead
# NOTE: python indexing is zero-based
print(sorted(pairs, key=lambda pair: pair[0]))

# You can learn more about other core data types and their methods here: 
# https://docs.python.org/3.7/library/stdtypes.html

# Because of its extensive standard library, Python is often described as coming with "batteries included".  
# Take a look at these "batteries": https://docs.python.org/3.7/library/

# You now know enough to complete this homework assignment (or at least where to look)
```

Output:
```
h
e
l
l
o
[2, 3, 4, 5]
[2]
hello, Josuke!
Howdy, partner!
13
Hi, Fred!
[('radical', 4), ('analysis', 7), ('bighorn', 12), ('bounce', 32)]
[('analysis', 7), ('bighorn', 12), ('bounce', 32), ('radical', 4)]
```

```python
from typing import Any, Callable, Dict, Iterable, List, Tuple, Union
from collections import Counter
from math import isclose
```

```python
# we'll be using this for some of our tests
def test_isclose(x: List[float], y: List[float], rel_tol=1.e-6) -> bool:
    if len(x) != len(y):
        return False
    for i in range(len(x)):
        if not isclose(x[i], y[i], rel_tol=rel_tol):
            return False
    return True
```

# Overview

In this assignment, you will write functions to ... 

- calculate the dot product

- calculate Euclidean distance

- normalize a vector

- calculate cosine similarity and cosine distance

- identify the centroid of a set of vectors

- identify the medoid of a set of vectors

# Vector math from scratch

You're going to be implementing all of the functions listed above in pure Python (no external dependencies).  By now you've seen the math and implementations in [NumPy (**num**erical **Py**thon)](https://numpy.org/).  While using NumPy would be far more efficient, this is an exercise geared toward deepening our understanding of these concepts.

# Warm-up

It's important to stretch before exercising to avoid injury.  Likewise, let's warm up our programming muscles by scaling some vectors.

## Scalar division

Implement the function `divide_by_scalar` which should divide each element in a list of floats by a scalar value.

<!-- grade_id: cell-f9fb6f774b34fc1e, points: 0, grade: false, solution: true -->
```python
def divide_by_scalar(x: List[float], scalar: int) -> List[float]:
    """
    Takes a list of floats representing the vector $x$ and divides each element by a scalar value.
    """
    # YOUR CODE HERE
    if scalar == 0:
        raise ValueError("Cannot divide by zero!")
    
    return [elem / scalar for elem in x]
```

<!-- grade_id: cell-d183ad85507a574c, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert divide_by_scalar([12., 3., 5.3], 4) == [3.0, 0.75, 1.325]
assert divide_by_scalar([0., 0., 0.], 17) == [0., 0., 0.]
```

# Normalizing vectors

Let's implement some methods to calculate distance from the origin and normalize vectors.

## `p_norm`

Implement the function `p_norm` to calculate the $p$-norm according to the following formula: 

$$
\vert \vert x \vert \vert{}_{p} = [\sum_{i=0}^{\vert x \vert} abs(x_{i})^{p}]^{\frac{1}{p}}
$$

This should take a list representing $\vec{x}$ and return the $p$-norm for $\vec{x}$.

<!-- grade_id: cell-f0b2a0dd285a4f1f, points: 0, grade: false, solution: true -->
```python
def p_norm(x: List[float], p: int) -> float:
    assert p >= 1
    """
    Takes a List (x) representing a vector and p.
    Returns the $p$-norm for $\vec{x}$.
    """
    # YOUR CODE HERE
    assert p >= 1, "p should be a positive integer greater than or equal to 1"
    
    return sum(abs(xi) ** p for xi in x) ** (1 / p)
```

<!-- grade_id: cell-1d80af31d2655fd1, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

# the L1 norm
assert isclose(p_norm([1., 2., 3.], p=1), 6., rel_tol=1.e-0)
```

<!-- grade_id: cell-c5fce0e9b484341a, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

# the L2 norm
assert isclose(p_norm([1., 2., 3.], p=2), 3.74165738, rel_tol=1.e-6)
```

<!-- grade_id: cell-bee99bf1e79f9d9e, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert isclose(p_norm([1., 2., 3.], p=3), 3.301927, rel_tol=1.e-6)
```

<!-- grade_id: cell-3d4cc66d6dd5b305, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert p_norm([1., 2., 3.], p=1) == 6
```

<!-- grade_id: cell-b3f9526f2acd13a4, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert isclose(p_norm([7., 14., 3.], p=2), 15.9373774, rel_tol=1.e-6)
```

## Vector normalization

The normalized form of a vector $\textbf{x}$ uses the 2-norm (Euclidean norm) to normalize each term in some vector $\mathbf{x}$:
    
    
$$
\begin{aligned}
\text{norm}(\mathbf{x}) &= \frac{\mathbf{x}}{\vert \vert  x \vert \vert_{2}}
\\[1em]
&= \frac{\mathbf{x}}{\left[\sum_{i=0}^{\vert \mathbf{x} \vert} abs(x_{i})^{2}\right]^{\frac{1}{2}}}
\\[1em]
&= \frac{\mathbf{x}}{\sqrt{\sum_{i=0}^{\vert \mathbf{x} \vert} abs(x_{i})^{2}}}
\end{aligned}
$$

<!-- grade_id: cell-ecadb414cfc35ee7, points: 0, grade: false, solution: true -->
```python
def normalize(x: List[float]) -> List[float]:
    """
    Returns the 2-norm of the vector x.
    """
    # YOUR CODE HERE
    norm = sum(xi ** 2 for xi in x) ** 0.5
    
    if norm == 0:
        return [0 for xi in x] 
    return [xi / norm for xi in x]
```

<!-- grade_id: cell-ca122481690b7a62, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert test_isclose(normalize([1., 2., 3.]), [0.26726124, 0.53452248, 0.80178373], rel_tol=1.e-6)
```

<!-- grade_id: cell-cd147030e714486d, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert test_isclose(normalize([10., 20., 30.]), [0.26726124, 0.53452248, 0.80178373], rel_tol=1.e-6)
```

<!-- grade_id: cell-0a05808454142996, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert test_isclose(normalize([100., 200., 300.]), [0.26726124, 0.53452248, 0.80178373], rel_tol=1.e-6)
```

<!-- grade_id: cell-664b4d476d5f6b80, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert test_isclose(normalize([1000., 2000., 3000.]), [0.26726124, 0.53452248, 0.80178373], rel_tol=1.e-6)
```

<!-- grade_id: cell-eac14c40e001d7b2, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert test_isclose(normalize([52., -2., 13.]), [0.96946785, -0.03728723, 0.24236696], rel_tol=1e-05)
```

<!-- grade_id: cell-88e2f559f73bff9c, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert test_isclose(normalize([520., -20., 130.]), [0.96946785, -0.03728723, 0.24236696], rel_tol=1e-05)
```

<!-- grade_id: cell-0c0d88e0df3c5110, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()


# HINT: There are several ways to avoid a divide-by-zero error; 
# you can check for zero values before dividing, 
# but another method is to use a *really* small number instead of zero

assert test_isclose(normalize([0., 0., 0.]), [0., 0., 0.], rel_tol=1e-02)
```

# Comparing vectors

Now that we can measure **magnitude** (L2 norm), we can explore measuring distance in terms magnitude, direction, or both.

## Dot product

Implement the function `dot_product`. This should take two lists representing $\vec{x}$ and $\vec{y}$ and return a single number.

Notes:
- the two lists must have the same number of dimensions 

<!-- grade_id: cell-61af8e7a542daed6, points: 0, grade: false, solution: true -->
```python
def dot_product(x: List[float], y: List[float]) -> float:
    """
    Returns the inner product of two 1D vectors.
    """
    assert len(x) == len(y)
    # YOUR CODE HERE
    return sum(xi * yi for xi, yi in zip(x, y))
```

<!-- grade_id: cell-9d480581d048f27f, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert dot_product([0., 0., 0.],   [0., 0., 0.]) == 0.
```

<!-- grade_id: cell-15a5a29d62ff60d6, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert dot_product([0., 2., 0.],   [0., 0., 0.]) == 0.
```

<!-- grade_id: cell-414e07a849804d0c, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert dot_product([0., 0., 0.],   [0., 2., 0.]) == 0.
```

<!-- grade_id: cell-4212376a50d617b3, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert dot_product([12., 3., 7.], [2., 4., 4.]) == 64.
```

## `euclidean_distance`

$\text{d}(\mathbf{a}, \mathbf{b}) = \sqrt{\sum_{i=0}^{\vert \mathbf{a} \vert} (a_{i} - b_{i})^{2}}$ where $\text{d}$ represent the Euclidean distance of vectors $\mathbf{a}$ and $\mathbf{b}$.

Implement the function `euclidean_distance`. This should take two lists representing $\vec{x}$ and $\vec{y}$ and return a single number.

<!-- grade_id: cell-33628eb13ae97d30, points: 0, grade: false, solution: true -->
```python
def euclidean_distance(x: List[float], y: List[float]) -> float:
    """
    Returns in the Euclidean distance of two 1D vectors.
    """
    # YOUR CODE HERE
    squared_diff_sum = sum((xi - yi) ** 2 for xi, yi in zip(x, y))
    
    return squared_diff_sum ** 0.5
```

<!-- grade_id: cell-d3378aeb39de8228, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert euclidean_distance([1., 2., 3.], [1., 2., 3.]) == 0.
```

<!-- grade_id: cell-cc121bbc40eab257, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert euclidean_distance([1., 2., 3.], [2., 2., 3.]) == 1.
```

<!-- grade_id: cell-b2829e2d37249348, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert euclidean_distance([1., 2., 3.], [0., 2., 3.]) == 1.
```

<!-- grade_id: cell-ae6dd818c861a073, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

# The Euclidean distance of some vector $v$ and the origin is equal to the L2 norm of $v$
assert euclidean_distance([1., 2., 3.], [0., 0., 0.]) == p_norm([1., 2., 3.], p=2)
```

## Cosine similarity

$$
\begin{aligned}
cos(\mathbf{a}, \mathbf{b}) &= \frac{\mathbf{a} \cdot \mathbf{b}}{\vert \vert a \vert \vert_{2}  \times \vert \vert b \vert \vert_{2}}
\end{aligned}
$$
Implement a method that calculates the cosine similarity of a pair of vectors.

<!-- grade_id: cell-dc629e60da20d82e, points: 0, grade: false, solution: true -->
```python
def cosine_similarity(x: List[float], y: List[float]) -> float:
    """
    Calculates the cosine similarity of two vectors.
    """
    # YOUR CODE HERE
    dot_prod = dot_product(x, y)
    norm_x = p_norm(x, 2)
    norm_y = p_norm(y, 2)
    
    return dot_prod / (norm_x * norm_y) if norm_x != 0 and norm_y != 0 else 0
```

<!-- grade_id: cell-ab5a6c75a2cab391, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert isclose(cosine_similarity([100., 200., 300.], [100., 200., 300.]), 1., rel_tol=1.e-6)
```

<!-- grade_id: cell-699c738b228c61e2, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert isclose(cosine_similarity([100., 200., 300.], [0.26726124, 0.53452248, 0.80178373]), 1., rel_tol=1.e-6) 
```

<!-- grade_id: cell-be32f4bf64d55ed1, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

assert isclose(cosine_similarity([2., 2., 2.], [-2., -2., -2.]), -1., rel_tol=1.e-2)
```

<!-- grade_id: cell-7b4261bcf4a898e0, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

# orthogonal vectors
assert isclose(cosine_similarity([2., 0.], [0., 2.]), 0., rel_tol=1.e-2)

# if cosine is 0, dot product is also zero
assert cosine_similarity([2., 0.], [0., 2.]) == dot_product([2., 0.], [0., 2.])
```

# Foundations of clustering

Now that you've implemented a few functions to caclulate distance and similarity, we can think about how to represent and compare groups of vectors.

## Centroids

$$
c_{x} = \frac{1}{\vert X \vert}\sum_{x_{i} \in X} x_{i} 
$$

where $X$ is a matrix and ${x}_{i}$ represents a row (vector) in $X$.

A **centroid** is an average or mean vector for some group of vectors.  Generally it does not correspond to an actual vector from the group, but rather one between them all (i.e., the one in the center).  Implement the function `find_centroid` to calculate the centroid for a group of vectors.

<!-- grade_id: cell-f8cdd63f8c2269c1, points: 0, grade: false, solution: true -->
```python
def find_centroid(X: List[List[float]]) -> List[float]:
    """
    Calculates centroid given a matrix X.
    X is represented using a list of lists.
    Each inner list is a vector.
    """
    # all vectors must have the same dimensionality
    assert len(set(len(v) for v in X)) == 1
    # YOUR CODE HERE
    num_vectors = len(X)
    num_dimensions = len(X[0])
    
    centroid = [0.0] * num_dimensions
    
    for vector in X:
        for i in range(num_dimensions):
            centroid[i] += vector[i]
    
    centroid = [val / num_vectors for val in centroid]
    
    return centroid
```

<!-- grade_id: cell-4164e7df7416851c, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

m = [
    [1, 1, 1],
    [5, 5, 5]
]

res = find_centroid(m)

assert test_isclose(res, [3.0, 3.0, 3.0], rel_tol=0)
```

<!-- grade_id: cell-0d9901d0b129bf37, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

m = [
    [1, 1, 1],
    [2, 2, 2],
    [5, 5, 5]
]

res = find_centroid(m)

assert test_isclose(res, [2.666, 2.666, 2.666], rel_tol=1.e-3)
```

## Medoids

$$
\mathbf{X}_{medoid} = \underset{v \in \mathbf{X}}{\text{argmin}}\frac{\sum_{x_{i} \in X} d(v, x_{i})}{\vert X \vert}
$$

$\ldots$ where 
 - $d$ is some distance function.
 - $\text{argmin}$ finds the $v$ with the minimum average distance.


Given a group of vectors $\mathbf{X}$, the **medoid** of $\mathbf{X}$ is the vector in $\mathbf{X}$ with the smallest average distance with every other vector in $\mathbf{X}$ according to some distance function $d$ (e.g. Euclidean distance).  _Note that this is not necessarily the vector in $\mathbf{X}$ nearest to the centroid._

Implement the function `find_mediod` to calculate the medoid for a group of vectors.

<!-- grade_id: cell-e8bafbad400cc929, points: 0, grade: false, solution: true -->
```python
# type for a function that takes a pair of params representing two vectors and returns a float
DistFunction = Callable[[List[float], List[float]], float]

def find_medoid(X: List[List[float]], dist_function: DistFunction) -> List[float]:
    """
    Finds the distance medoid of X (i.e., the vector in X that is on average closest to every other vector)
    """
    # YOUR CODE HERE
    best_medoid = None
    best_average_distance = float('inf')
    
    for candidate in X:
        total_distance = sum(dist_function(candidate, other) for other in X if other != candidate)
        average_distance = total_distance / (len(X) - 1)
        
        if average_distance < best_average_distance:
            best_average_distance = average_distance
            best_medoid = candidate
    
    return best_medoid
```

<!-- grade_id: cell-736e9d9aeb2d987d, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

m = [
    [1, 1, 1],
    [2, 2, 2],
    [5, 5, 5]
]

res = find_medoid(m, euclidean_distance)

assert test_isclose(res, [2., 2., 2.], rel_tol=1.e-3)
```

<!-- grade_id: cell-66865d9dea79aa33, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

m = [
    [1, 1, 1],
    [10, 20, 20],
    [5, 5, 5]
]

# cosine distance
cdist = lambda x, y: 1 - cosine_similarity(x, y)
res = find_medoid(m, cdist)
assert test_isclose(res, [1., 1., 1.], rel_tol=1.e-3)
```

<!-- grade_id: cell-128e86671029cfce, points: 1, grade: true, solution: false -->
```python
# verify the NumPy module is not in scope
assert "numpy" not in dir()

m = [
    [3, 5, 1],
    [0, 0, 0],
    [3, 5, 1],
    [3, 5, 1],
    [2, 1, 0],
    [2, 1, 0],
    [2, 1, 0],
    [2, 1, 1],
    [0, 0, 1],
    [-1, -1, 0],
    [2, 2, 4]
]

# cosine distance
cdist = lambda x, y: 1 - cosine_similarity(x, y)
res = find_medoid(m, cdist)
assert test_isclose(res, [2., 1., 1.], rel_tol=1.e-3)
```

If you've successfully implemented all of these functions, then you now know enough to perform unsupervised learning (clustering).