## Tips
- To avoid unpleasant surprises, I suggest you _run all cells in their order of appearance_ (__Cell__ $\rightarrow$ __Run All__).


- If the changes you've made to your solution don't seem to be showing up, try running __Kernel__ $\rightarrow$ __Restart & Run All__ from the menu.


- Before submitting your assignment, make sure everything runs as expected. First, restart the kernel (from the menu, select __Kernel__ $\rightarrow$ __Restart__) and then **run all cells** (from the menu, select __Cell__ $\rightarrow$ __Run All__).

## Reminder

- Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name, UA email, and collaborators below:



Several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

```python
NAME = "Kathleen Costa"
# University of Arizona email address
EMAIL = "kathleencosta@arizona.edu"
# Names of any collaborators.  Write N/A if none.
COLLABORATORS = "N/A"
```

## Scratchpad

You are welcome to create new cells (see the __Cell__ menu) to experiment and debug your solution.

```python
%load_ext autoreload
%autoreload 2
```

# Mini Python tutorial

This course uses Python 3.11.

Below is a very basic (and incomplete) overview of the Python language... 

For those completely new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.11/library/stdtypes.html#common-sequence-operations).

```python
# This is a comment.  
# Any line starting with # will be interpreted as a comment

# this is a string assigned to a variable
greeting = "hello"

# If enclosed in triple quotes, strings can also be multiline:

"""
I'm a multiline
string.
"""

# let's use a for loop to print it letter by letter
for letter in greeting:
    print(letter)
    
# Did you notice the indentation there?  Whitespace matters in Python!

# here's a list of integers

numbers = [1, 2, 3, 4]

# let's add one to each number using a list comprehension
# and assign the result to a variable called res
# list comprehensions are used widely in Python (they're very Pythonic!)

res = [num + 1 for num in numbers]

# let's confirm that it worked
print(res)

# now let's try spicing things up using a conditional to filter out all values greater than or equal to 3...
print([num for num in res if not num >= 3])

# Python 3.7 introduced "f-strings" as a convenient way of formatting strings using templates
# For example ...
name = "Josuke"

print(f"{greeting}, {name}!")

# f-strings are f-ing convenient!


# let's look at defining functions in Python..

def greet(name):
    print(f"Howdy, {name}!")

# here's how we call it...

greet("partner")

# let's add a description of the function...

def greet(name):
    """
    Prints a greeting given some name.
    
    :param name: the name to be addressed in the greeting
    :type name: str
    
    """
    print(f"Howdy, {name}!")
    
# I encourage you to use docstrings!

# Python introduced support for optional type hints in v3.5.
# You can read more aobut this feature here: https://docs.python.org/3.7/library/typing.html
# let's give it a try...
def add_six(num: int) -> int:
    return num + 6

# this should print 13
print(add_six(7))

# Python also has "anonymous functions" (also known as "lambda" functions)
# take a look at the following code:

greet_alt = lambda name: print(f"Hi, {name}!")

greet_alt("Fred")

# lambda functions are often passed to other functions
# For example, they can be used to specify how a sequence should be sorted
# let's sort a list of pairs by their second element
pairs = [("bounce", 32), ("bighorn", 12), ("radical", 4), ("analysis", 7)]
# -1 is last thing in some sequence, -2 is the second to last thing in some seq, etc.
print(sorted(pairs, key=lambda pair: pair[-1]))

# we can sort it by the first element instead
# NOTE: python indexing is zero-based
print(sorted(pairs, key=lambda pair: pair[0]))

# You can learn more about other core data types and their methods here: 
# https://docs.python.org/3.7/library/stdtypes.html

# Because of its extensive standard library, Python is often described as coming with "batteries included".  
# Take a look at these "batteries": https://docs.python.org/3.7/library/

# You now know enough to complete this homework assignment (or at least where to look)
```

# Introduce yourself

Congratulations on successfully launching the notebook!  That was more than half of the challenge for this assignment. 

Please briefly introduce yourself below.  

- How would you like your instructor to refer to you?
- Where are you from?  
- What do you hope to learn in this class?  
- What are you career goals?
- Anything else you'd like to share?

<!-- grade_id: greetings, points: 0, grade: true, solution: true -->
Hello Dr. Hahn-Powell! My name is Kathleen Costa, and I currently live in Atlanta, GA. I'm honestly hoping to get a concrete foundation of human language technology and all the programming languages associated (there are so many to learn all at once!). My career goals are to hopefully, one day, create my own language learning app that focuses on task-based language learning.

NOTE: several of the cells in this notebook are **read only** to ensure instructions aren't unintentionally altered.  

If you can't edit the cell, it is probably intentional.

# Overview

Congratulations on launching your first assignment!  The problems below just test that you're able to do some basic programming in Python.  Remember that basic programming proficiency (in any language) is a prerequisite of this course.  

# Implement `multiply(a, b)`

To test your understanding of Python basics, implement the following function which should return the product of its parameters `a` and `b`.  Add you solution after the comment **YOUR CODE HERE** and remove the `raise NotImplementedError`.  

### NOTES 

- _Don't overthink this. It's meant to be easy!_

- Usage of `: int` and `-> int` may be unfamiliar to you. These are [**type annotations**](https://www.python.org/dev/peps/pep-0483/), an optional feature supported by Python versions $\geq 3.5.\textrm{x}$.  
  - `: int` denotes that the data type is an integer 
  - `int`.  `-> int` denotes that the **return type** is an integer.  
  - Type hints can help you understand expectations for program input, output, and how the input might be manipulated.  Type information can be used to check for certain errors.  For instance, You can't calculate [`len()`](https://docs.python.org/3/library/functions.html#len) (length) of an `int`, but you can calculate [`len()`](https://docs.python.org/3/library/functions.html#len) of a string.

<!-- grade_id: multiply-problem, points: 0, grade: false, solution: true -->
```python
def multiply(a: int, b: int) -> int:
    """
    Takes two integers and returns their product
    """
    # YOUR CODE HERE
    return a*b
```

<!-- grade_id: multiply-tests, points: 3, grade: true, solution: false -->
```python
assert multiply(2, 2) == 4
assert multiply(2, 0) == 0
assert multiply(-1, -1) == 1
```

# Reverse a list and remove the first element
This is exercise is just to illustrate how you can use the test cases to inform your implementation.

Implement the function `weird_request(seq)` which ... 
1. takes a list
2. reverses its contents
3. drops the element at the front
4. and returns what remains

For those new to Python, [this section of the official documentation may be useful](https://docs.python.org/3.7/library/stdtypes.html#common-sequence-operations)

<!-- grade_id: weird-request-problem, points: 0, grade: false, solution: true -->
```python
def weird_request(seq):
    """
    :param seq: a list of anything
    takes a list, reverses its contents, drops the element at the front, and returns what remains.
    """
    # YOUR CODE HERE
    seq.reverse
    
    reverse_seq = seq[::-1]
    return reverse_seq [1:]
```

<!-- grade_id: cell-7354f1ffb848a795, points: 2, grade: true, solution: false -->
```python
assert weird_request(["a","b","c"]) == ["b","a"]
assert weird_request([1,2,3]) == [2,1]
```

<!-- grade_id: cell-5a21d789e65281cf, points: 2, grade: true, solution: false -->
```python
assert weird_request([1]) == []
assert weird_request([]) == []
```